## Week 3: Probability

### Version 1.1



Plese enter your name =""

## 1. Probability

<strong>1.1</strong> \[1 pt\] A college is looking to form a team of 8 to participate in a coding compeition. The team needs at least 2 females and 2 males. If there are 4 males and 7 females who want to join the competition, how many different ways can the team be formed? Please store your answer as an integer (e.g. ANS1 = 1) into the variable <strong>ANS1</strong>.

<!-- grade_id: cell-ee2bd7ab2781aa26, points: 0, grade: false, solution: true -->
```python
import math
import numpy as np
import numbers

# ffffffmm, fffffmmm, ffffmmmm

two_males = math.comb(4,2) * math.comb(7,6)
three_males = math.comb(4,3) * math.comb(7,5)
four_males = math.comb(4,4) * math.comb(7,4)


ANS1= two_males + three_males + four_males
```

<!-- grade_id: cell-2380314f88f56d22, points: 1, grade: true, solution: false -->
```python
assert type(ANS1) == int, "Problem 1.1, testing ANS1, type of value stored in variable does not match the expected type. Expecting Integer"
```

## 2. Independent Probability 

Tom and Lily are trying to solve a problem. The probability of Tom successfully solving the problem is 0.4 and the probability of Lily successfully solving the problem is 0.5. They do not influence each other when they are solving the problem. For each solution, store the value as accurately as you can (ex. ANS = 2/9 instead of .2222).

<strong>2.1</strong> \[1 pt\] What is the probability that both of them sucessfully solve the problem? Please store your answer to the variable <strong>ANS21</strong>.

<!-- grade_id: cell-a89dfc965132afad, points: 0, grade: false, solution: true -->
```python
# P(a and b) =
a = .4
b = .5
result = (a*b)
ANS21 = result
```

<!-- grade_id: cell-4d4bf277ef89e5a1, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS21, numbers.Number), "Problem 2.1, testing ANS21, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>2.2</strong> \[1 pt\] What is the probability that **exactly** one of them successfully solves the problem? Please store your answer to the variable <strong>ANS22</strong>. 

<!-- grade_id: cell-e55df99dc649e767, points: 0, grade: false, solution: true -->
```python
# p(x=1)
a = .4
b = .5
result = (a*(1-b))+((1-a)*b)
ANS22 = result
```

<!-- grade_id: cell-04bf980d4b26a9f0, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS22, numbers.Number), "Problem 2.2, testing ANS22, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>2.3</strong> \[1 pt\] Given that **exactly** one of them successfully solved the problem, what is the probability that the person is Tom?
Please store your answer to the variable <strong>ANS23</strong>. 

<!-- grade_id: cell-9cd919f33acab2ef, points: 0, grade: false, solution: true -->
```python
# P( Tom |xr=1)
a = .4
b = .5

ANS23= (a*(1-b))/b
```

<!-- grade_id: cell-9beab75a7dab793d, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS23, numbers.Number), "Problem 2.3, testing ANS23, type of value stored in variable does not match the expected type. Expecting Float."
```

## 3. Conditional Probability

There are two tests for a disease, one is rapid and the other is slow.  Given that an individual is infected, the rapid test will register positive 40% of the time, while the slow test will register positive 80% of the time; additionally, both tests will be positive 35% of the time.   

<strong>3.1</strong> \[1 pt\] Given an infected person and that the rapid test measures a positive results, what is the probability that the slow test is also positive?  Store your answer in the variable <strong>ANS31</strong>.  For grading (on this and future problems) store the answer as accurately as possible using standard python data types.  For example, if the answer is 1/3, set ANS31 = 1/3 instead of .3333.



<!-- grade_id: cell-8566435eee91457e, points: 0, grade: false, solution: true -->
```python
#P(slow positive | rapid test positive)
rapid_positive = .4
both_positive = .35
slow_positive = .8

ANS31 = both_positive/rapid_positive
```

<!-- grade_id: cell-557a697b7229071a, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS31, numbers.Number), "Problem 3.1, testing ANS31, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>3.2</strong> \[1 pt\]  Given an infected person and that the rapid test measures a negative result, what is the chance that the slow test is positive.  Store you answer in <strong>ANS32</strong>.  

<!-- grade_id: cell-d8a6247556748dcc, points: 0, grade: false, solution: true -->
```python
#p(slow positive | 1-rapid positive) = p(1-rapid positive & solow positive)/1-rapid positive
rapid_positive = .4
both_positive = .35
slow_positive = .8
result = (slow_positive - both_positive)/(1-rapid_positive)
ANS32 = result
```

<!-- grade_id: cell-0c3c06d832ee2f42, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS32, numbers.Number), "Problem 3.2, testing ANS32, type of value stored in variable does not match the expected type. Expecting Float."
```

## 4. Bayes Rule

Suppose in the above example that people not infected always test negative for both tests.

<strong>4.1</strong> \[1 pt\] Assume currently, 20% of the population has the virus.  Use Bayes theorem to calculate the chance that a persons has the virus conditioned on getting negative results for both tests?  Store your answer in the variable <strong>ANS41</strong> as accurately as possible.



<!-- grade_id: cell-f68142d6a4ebf595, points: 0, grade: false, solution: true -->
```python
#P( negitive rapid && negitive slow |has virus)
rapid_positive = .4 # rapid positive | virus
both_positive = .35 # rapid positive && slow positive | virus
slow_positive = .8 # slow postive | virus
has_virus = .20
no_viris = 1 - has_virus
c_both_positive = 1 -(rapid_positive + slow_positive - both_positive)
bayes = (has_virus*c_both_positive)/((has_virus*c_both_positive) + slow_positive)

ANS41 = bayes
```

<!-- grade_id: cell-c32efd59d0bc94f7, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS41, numbers.Number), "Problem 4.1, testing ANS41, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>4.2</strong> \[1 pt\]  Of the people in the population who are tested, 75% of their results from the slow test are positive.  What is the chance that a persons has the virus conditioned on getting negative result on the slow test?
Store your answer in the variable <strong>ANS42</strong> as accurately as possible.

<!-- grade_id: cell-2583cf0d6409ae0f, points: 0, grade: false, solution: true -->
```python
#P(negative slow | has virus)
a=.75
ANS42 = ((A*slow_positive)*(1-a))/(1-slow_positive)
```

Output:
```
0.5625
```

<!-- grade_id: cell-2cd6eebb1477a722, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS42, numbers.Number), "Problem 4.2, testing ANS42, type of value stored in variable does not match the expected type. Expecting Float."
```

Suppose that you are an art dealer for Van Gogh paintings. Whenever someone brings in a painting that they claim is from Van Gogh, you send in the painting for pigment analysis. Pigment analysis successfully identifies 90% of Van Gogh paintings and falsely identifies 20% of forgeries as being from Van Gogh. In your experience, 99% of the paintings that people bring to you are forgeries.

<strong>4.3</strong> \[1 pt\] Use Bayes theorem to calculate the chance that a painting is from Van Gogh if the pigment analysis finds that it is from Van Gogh. Store your answer in the variable <strong>ANS43</strong> as accurately as possible.

<!-- grade_id: cell-fc6f0a9fc9408cc2, points: 0, grade: false, solution: true -->
```python
fake = .99
real = 1 - fake
#p(pigment say real | its fake)
pigment_real_is_fake = .2
pigment_real_is_real = .9
#p(is real | pigment real)
result = (pigment_real_is_real*real)/((pigment_real_is_real*real)+(pigment_real_is_fake*fake))

ANS43 = result

```

<!-- grade_id: cell-f3787978ae55118b, points: 1, grade: true, solution: false -->
```python
assert isinstance(ANS43, numbers.Number), "Problem 4.3, testing ANS43, type of value stored in variable does not match the expected type. Expecting Float."
```

## 5. Expectation and Variance

In a state park, a lost person arrives at a fork with 3 roads: <br>
• The first road leads back to the same point after 1 hour of walking.<br>
• The second road leads back to the same point after 2 hours of walking.<br>
• The third road leads to the parking lot after 5 hours of walking. <br>
There are no signs on the roads and when the lost person arrives back at the fork they have no memory of which road(s) they had taken previously.

<strong>5.1</strong> \[2 pt\]  Assuming that the lost person always chooses each road with equal probability, what is the expected amount of time until the tourist arrives at the parking lot?  Store your answer in the variable <strong>ANS51</strong> as accurately as possible.

Hint: If you are stuck on how to start, go back and watch the lectures on expectation where we compute a similar example.

<!-- grade_id: cell-f68142d6a4ebf596, points: 0, grade: false, solution: true -->
```python
import math

ANS51 = math.comb(3,1)+ math.comb(3,2) + math.comb(3,3)+ math.comb(3,0)

```

Output:
```
8
```

<!-- grade_id: cell-c32efd59d0bc94f3, points: 2, grade: true, solution: false -->
```python
assert isinstance(ANS51, numbers.Number), "Problem 5.1, testing ANS51, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>5.2</strong> \[2 pt\]  What is the expectation and variance of the sum of the roll of two dice?   Store your answer as a tuple in the variable <strong>ANS52</strong> as accurately as possible.

<!-- grade_id: cell-bbd9380cfee320a3, points: 0, grade: false, solution: true -->
```python
import statistics
n = np.array([1,2,3,4,5,6])
p_of_n = n*(1/6)
v = (n -sum(p_of_n))**2 *1/6

ANS52=(2*sum(p_of_n),2*sum(v))

```

Output:
```
(7.0, 5.833333333333334)
```

<!-- grade_id: cell-402ceeaf0297d261, points: 2, grade: true, solution: false -->
```python
assert len(ANS52) == 2,  "Problem 5.2, testing ANS52, type of value stored in variable does not match the expected type. Tuple." 
assert isinstance(ANS52[0], numbers.Number) and isinstance(ANS52[1], numbers.Number) , "Problem 5.2, testing ANS52, type of values stored in tuples do not match the expected type. Expecting Float."
```

## 6. Monte Carlo Method



<strong>6.1</strong> \[2 pt\]  You receive one of five distinct coupons uniformily at random each day.  Use the Monte Carlo Method to estimate to within .01 days, the expected number of days until you collect two copies of each of the five coupons.  Record your answer in <strong>ANS61</strong>.

Hint: If your code is not very efficient, the autograder may timeout.  You can run this on your own computer and then copy the answer to this notebook. 

<!-- grade_id: cell-ceb4705829e49a2c, points: 0, grade: false, solution: true -->
```python
import numpy as np

n=100000
high_range=5
num_of_days=0

def check(coupon_list):
    for i in range(high_range):
        if coupon_list[i] < 2:
            return False
    return True

for _ in range(n):
    day = 1
    coupons = [0 for i in range(high_range)]
    while True:
        random_coupon = np.random.randint(0,high_range)
        coupons[random_coupon] = coupons[random_coupon] + 1
        if check(coupons):
            break
        else:
            day+=1
    
    num_of_days+=day

ANS61 = round(num_of_days/n,2)

```

<!-- grade_id: cell-24a0cb9f1d006a9b, points: 2, grade: true, solution: false -->
```python
assert isinstance(ANS61, numbers.Number), "Problem 6.1, testing ANS61, type of value stored in variable does not match the expected type. Expecting Float."
```

<strong>6.2</strong> \[2 pt\]  There are 30 coins.  While 29 of them are fair, 1 of them flips heads with probability 60%.  You flip each coin 100 times and record the number of times that it lands heads.  You then order the coins from most heads to least heads.  You seperate out the 10 coins that flipped heads the most into a pile of "candidate coins".  If several coins are tied for the 10th most heads, include them all.  (So your pile of candidate coins will always contain at least 10 heads, but may also include more).     

Use the Monte Carlo method to compute (within .1%) the probability that the unfair coin is in the pile of candidate coins.   Record your answer in <strong>ANS62</strong>.

Hint 1: use np.random.binomial to speed up simulation.  A binomial variable with parameters n and p is the number of heads resulting from flipping n coins, where each has probability p of landing heads.  

Hint 2: If your code is not very efficient, the autograder may timeout.  You can run this on your own computer and then copy the answer.  

<!-- grade_id: cell-49d189d155b7bffe, points: 0, grade: false, solution: true -->
```python
fair = np.average(np.random.binomial(100,.5, 30))
bad = np.average(np.random.binomial(100,.6, 30))
ANS62 = (bad*.01)/(fair*.99)+(bad*.01)

```

<!-- grade_id: cell-691ed4f14f4a3360, points: 2, grade: true, solution: false -->
```python
assert isinstance(ANS62, numbers.Number), "Problem 6.2, testing ANS62, type of value stored in variable does not match the expected type. Expecting Float."
```